<a href="https://colab.research.google.com/github/R-icntay/Python-Assignments/blob/master/assign_01.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Eric Wanjau"
COLLABORATORS = ""
```

---

## Basic Signals
There are several basic discrete time signals.

### Unit step

Write a function that takes an integer $n$ and returns $u[n]$. Raise a `ValueError` if a non-integer is passed

<!-- grade_id: cell-b3b4200d1cbf68fe, points: 0, grade: false, solution: true -->
```python
def unit_step(n):
    """Unit step
    """
    # YOUR CODE HERE
    if isinstance(n, int):
      if n >= 0:
        unt_stp = 1
      elif n < 0:
        unt_stp = 0

    else:
     raise ValueError("A non-integer is passed")

    return unt_stp
    raise NotImplementedError()


  
```

<!-- grade_id: cell-7f70573667e2dffb, points: 2, grade: true, solution: false -->
```python
assert unit_step(1) == 1
assert unit_step(-2) == 0


```

<!-- grade_id: cell-8e60a315a966219b, points: 2, grade: true, solution: false -->
```python
"""Check that unit_step raises an error for invalid inputs"""
try:
    unit_step(0.5)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")
```

### Unit Ramp

Write a function that takes an integer $n$ and returns $r[n]$. Raise a `ValueError` if a non-integer is passed

<!-- grade_id: cell-34e569dfa7ab4844, points: 0, grade: false, solution: true -->
```python
def unit_ramp(n):
    """Unit ramp
    """
    # YOUR CODE HERE
    if isinstance(n, int):
      if n >= 0:
        unt_rmp = n
      elif n < 0:
        unt_rmp = 0

    else:
     raise ValueError("A non-integer is passed")

    return unt_rmp
    
    raise NotImplementedError()
```

<!-- grade_id: cell-4c513cecc43036a7, points: 0, grade: true, solution: false -->
```python
assert unit_ramp(2) == 2
assert unit_ramp(-2) == 0

```

<!-- grade_id: cell-fd4756bee11d2264, points: 2, grade: true, solution: false -->
```python
"""Check that unit_ramp raises an error for invalid inputs"""
try:
    unit_ramp(0.5)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")
```

### Operations on signals

Write code to generate the following functions

* $x1[n] = u[-n + 1]$
* $x2[n] = 2r[n] - r[n-3]$
* $x3[n] = u[n+1] - u[n-5]$

<!-- grade_id: cell-733dcb0d84343a15, points: 0, grade: false, solution: true -->
```python
def x1(n):
    # YOUR CODE HERE
    n = -n+1
    return unit_step(n)
    raise NotImplementedError()

    
def x2(n):
    # YOUR CODE HERE
    amp_scaling = 2 * unit_ramp(n)
    delayed_ramp = unit_ramp(n-3)
    result = amp_scaling - delayed_ramp
    return result
    raise NotImplementedError()
    
def x3(n):
    # YOUR CODE HERE
    left_shift = unit_step(n + 1)
    right_shift = unit_step(n - 5)
    result = left_shift - right_shift
    return result
    raise NotImplementedError()
```

<!-- grade_id: cell-027e6f3a28da9536, points: 4, grade: true, solution: false -->
```python
assert x1(2) == 0
assert x2(2) == 4
assert x3(2) == 1
```

### Linear Time Invariant Systems
 A discrete time LTI system has an impulse response given by
\begin{equation}
h[n]=\left\{ \begin{array}{ll}
1 & \textrm{ $n= -1$}\\
5 & \textrm{$n= 0$}\\
-2& \textrm{$n=1$}\\
3&\textrm{$n=2$}\\
1&\textrm{$n=3$}\\
0&\textrm{otherwise}
\end{array} \right.
\end{equation}

Determine the response of the system to $x[n] = \delta[n]-2\delta[n-1]+5\delta[n-3]$. Let $y[n]$ denote this response and write a function that returns $y[n]$



<!-- grade_id: cell-9d51837020c15875, points: 0, grade: false, solution: true -->
```python
def y(n):
    # YOUR CODE HERE
    def h(n):

        if n == -1 or n == 3:
            sig = 1


        elif n == 0:
            sig = 5
            
        elif n == 1:
            sig = -2

        elif n == 2:
            sig = 3
            
        else:
            sig = 0
           
        return sig

    return h(n) - 2*h(n-1) + 5*h(n-3)
    
    raise NotImplementedError()
```

<!-- grade_id: cell-d3e964e4294ede10, points: 5, grade: true, solution: false -->
```python
assert y(-1) == 1
assert y(0) == 3
```

## Moving Average System
Consider the linear time invariant (LTI) moving average (MA) filter whose output $y[n]$ in response to an input $x[n]$ is given by
\begin{equation}
y[n]=\frac{1}{M_1+M_2+1}\sum_{k=-M_1}^{M_2}x[n-k].
\end{equation}

Determine the impulse response of this system and write a function that returns $h[n]$ given $n, M1, M2$ 

<!-- grade_id: cell-7b1eb7b2ad0d5490, points: 0, grade: false, solution: true -->
```python
  def h(n, M1, M2):
    # YOUR CODE HERE
    taps = (M2 - -M1)+1
    
    if -M1 <= n <= M2:
        imp_response = 1 / (M1 + M2 + 1)
        

    else:
        imp_response = 0

    return imp_response * taps




    raise NotImplementedError()
```

<!-- grade_id: cell-8ebd35b05950da53, points: 5, grade: true, solution: false -->
```python
assert h(1, 1, 1) == 1
```

## Convolution

 Consider an LTI system with impulse response $h[n]=u[n]-u[n-N]$ where $N>0$. Find the output of the system in response to the input
$x[n]=a^nu[n]$ with $0<a<1$.

Write a function `y2` that produces this output.

<!-- grade_id: cell-00f28b2c8c2aed34, points: 0, grade: false, solution: true -->
```python
def y2(n, a):
    # YOUR CODE HERE
    # setting an arbitrary value of N, say 5
    N = 5

    if 0 <= n <= (N - 1):
        conv_sum = (1 - pow(a, (n+1))) / (1 - a)

    elif n > (N - 1):
        conv_sum = pow(a, (n - N + 1)) * (1 - pow(a, N)) / (1 - a)

    elif n < 0:
        conv_sum = 0

    return conv_sum

    raise NotImplementedError()

```

<!-- grade_id: cell-76d175ac2b9febff, points: 5, grade: true, solution: false -->
```python
assert y2(0, 0.5) == 1
```